```python
# Initialize Otter
import otter
grader = otter.Notebook("hw2.ipynb")
```

# FILL IN YOUR NAME AND THE NAME OF YOUR PEER (IF ANY) BELOW

**Name**: \<replace this with your name\>

**Peer**: \<replace this with your peer's name\>

## Collaboration policy
Students are responsible for writing their own quizzes, assignments, and exams. For homework assignments, students are welcome (and encouraged) to discuss problems with one peer, **but each student must write their own assignment wrtieup and code individually**. The peer must be listed at the top of the writeup for each assignment. *Note: I will treat AI assistants as peers. That is, students are welcome to discuss problems with an AI assistant, but it is considered cheating to directly obtain an answer by querying the assistant. Please credit any AI assistant that you use.*

# Homework 2 -- Dijkstra's Algorithm and A* (100 pts)

**Due:** Tuesday, February 11th, 2025 at 11:59 pm

This homework builds on the material in the slides and the LaValle chapter on Discrete Optimal Search.

We will use Jupyter/Colab notebooks throughout the semester for writing code and generating assignment outputs. 

## 1) Tracing Dijkstra and A* (60 pts)

Consider the following 3x3 grid where nodes are gonnected to their vartical, horizontal, and diagonal neighbors (N, E, S, W, NE, NW, SE, SW). The start node is (0,0) and the goal node is (2,2). The cost for vertical and horizontal actions is 1, and the cost of diagonal actions is 1.4, with two exceptions. There are conveyor belts between nodes `(1,0)` and `(1,1)`, and between nodes `(1,2)` and `(2,2)`, reducing the cost of those actions to 0.2. The grid with associated action costs is shown below.

<img src="https://jorge-a-mendez.github.io/images/ese564s24/dijkstragrid.png" width=50%>

**Throughout this question, assume that ties are broken by sorting based on the nodes' first coordinate and further ties are broken by the second coordinate (e.g., (0,1) comes before (1,0); (0,2) comes before (1,1); (0,1) comes before (0,2)).**


### 1.1) Dijkstra (16 pts)

We are going to manually trace Dijkstra's algorithm and list the order in which nodes are visited. We will start from a prioirity queue that contains a single element, the starting node:
```
q = [(0,0): 0]
```
And a parent list that contains `None` for the start state's parent:
```
p = {(0,0): None}
```
Each question below will ask you to add the current node's neighbors to the queue and add the parent of each new node to the `p` dictionary. *Hint: you may find it helpful to also track the `dist` map defined in the pseudocode in the slides.*



#### 1.1.1)

In the first step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the first iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_111 = ...
```

After the first step, what are the contents of the queue (i.e., after extracting the first node `v_111` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_111 = ...
```

After the first step, what are the contents of the parents directory (i.e., after extracting the first node `v_111` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_111 = ...
```

#### 1.1.2)

In the second step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the second iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_112 = ...
```

After the second step, what are the contents of the queue (i.e., after extracting the second node `v_112` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_112 = ...
```

After the second step, what are the contents of the parents directory (i.e., after extracting the second node `v_112` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_112 = ...
```

#### 1.1.3)

In the third step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the third iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_113 = ...
```

After the third step, what are the contents of the queue (i.e., after extracting the third node `v_113` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_113 = ...
```

After the third step, what are the contents of the parents directory (i.e., after extracting the third node `v_113` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_113 = ...
```

#### 1.1.4)

In the fourth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the fourth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_114 = ...
```

After the fourth step, what are the contents of the queue (i.e., after extracting the fourth node `v_114` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_114 = ...
```

After the fourth step, what are the contents of the parents directory (i.e., after extracting the fourth node `v_114` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_114 = ...
```

#### 1.1.5)

In the fifth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the fifth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_115 = ...
```

After the fifth step, what are the contents of the queue (i.e., after extracting the fifth node `v_115` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_115 = ...
```

After the fifth step, what are the contents of the parents directory (i.e., after extracting the fifth node `v_115` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_115 = ...
```

#### 1.1.6)

In the sixth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the sixth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_116 = ...
```

After the sixth step, what are the contents of the queue (i.e., after extracting the sixth node `v_116` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_116 = ...
```

After the sixth step, what are the contents of the parents directory (i.e., after extracting the sixth node `v_116` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_116 = ...
```

#### 1.1.7)

In the seventh step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the seventh iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_117 = ...
```

After the seventh step, what are the contents of the queue (i.e., after extracting the seventh node `v_117` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_117 = ...
```

After the seventh step, what are the contents of the parents directory (i.e., after extracting the seventh node `v_117` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_117 = ...
```

#### 1.1.8)

In the eihgth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the eighth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_118 = ...
```

After the eighth step, what are the contents of the queue (i.e., after extracting the eighth node `v_118` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_118 = ...
```

After the eihgth step, what are the contents of the parents directory (i.e., after extracting the eighth node `v_118` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_118 = ...
```

#### 1.1.9)

In the ninth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the ninth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_119 = ...
```

After the ninth step, what are the contents of the queue (i.e., after extracting the ninth node `v_119` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_119 = ...
```

After the ninth step, what are the contents of the parents directory (i.e., after extracting the ninth node `v_119` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_119 = ...
```

#### 1.1.10)

When Dijkstra succeeds, we can find a plan by tracing back the parents in the `p` dictionary. In particular, we find the goal node's parent, then the parent's parent, and so on, until we reach the start state, and then we revert this path. What is the path found by Dijkstra for this problem?

_Points:_ 1.5

```python
# Express your answer as a list of tuples starting from the start state and ending in the goal state. If no plan is found, set the plan to []
plan_1110 = ...
```

#### 1.1.11)

What is the total cost of the plan found by Dijstra's algorithm?

_Points:_ 1

```python
# Express your answer as a scalar
ans_1111 = ...
```


### 1.2) A* with Octile Heuristic (13 pts)


To use A*, we need to define a heuristic function that predicts the cost to reach the goal from each node. We will start by using the *octile heuristic*. The octile heuristic is defined as the minimum distance from a node to the goal assuming that horizontal and vertical actions have a cost of 1 and diagonal 45-degree actions have a cost of 1.4. More precisely, the heuristic is defined as $\max(\Delta x, \Delta y) + 0.4 \cdot \min(\Delta x, \Delta y)$, where $\Delta x$ and $\Delta y$ are the horizontal and vertical distance to goal, respectively.



Assuming that $\Delta x$ and $\Delta y$ are the horizontal and vertical distance to goal, respectively, which of the following formulas represents the octile heuristic?

a. $\max(\Delta x, \Delta y) + 1.4 \cdot \min(\Delta x, \Delta y)$\
b. $\max(\Delta x, \Delta y) + 0.4 \cdot \min(\Delta x, \Delta y)$\
c. $\Delta x + \Delta y + 0.4 \cdot \min(\Delta x, \Delta y)$\
d. $\max(\Delta x, \Delta y) + 0.4 \cdot (\Delta x + \Delta y)$

_Points:_ 3

```python
# Enter your answer as a lower-case Python string (e.g., "a")
ans_12 = ...
```

We are going to manually trace the A* algorithm using the octile heuristic, and list the order in which nodes are visited. We will start from a priority queue that contains a single element, the starting node, associated to the heuristic estimate of the cost to reach the goal:
```
q = {(0,0): 2.8}
```
And a parent list that contains `None` for the start state's parent:
```
p = {(0,0): None}
```
Each question below will ask you to add the current node's neighbors to the queue and add the parent of each new node to the `p` dictionary.

*Hint: in addition to `q` and `p`, it may be helpful for you to keep track of `c2come` and `ctot` as defined in the slides.*


#### 1.2.1)

In the first step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the first iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_121 = ...
```

After the first step, what are the contents of the queue (i.e., after extracting the first node `v_121` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_121 = ...
```

After the first step, what are the contents of the parents directory (i.e., after extracting the first node `v_121` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_121 = ...
```

#### 1.2.2)

In the second step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the second iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_122 = ...
```

After the second step, what are the contents of the queue (i.e., after extracting the second node `v_122` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_122 = ...
```

After the second step, what are the contents of the parents directory (i.e., after extracting the second node `v_122` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_122 = ...
```

#### 1.2.3)

In the third step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the third iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_123 = ...
```

After the third step, what are the contents of the queue (i.e., after extracting the third node `v_123` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_123 = ...
```

After the third step, what are the contents of the parents directory (i.e., after extracting the third node `v_113` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_123 = ...
```

#### 1.2.4)

Like Dijkstra, when A* succeeds, we can find a plan by tracing back the parents in the `p` dictionary. In particular, we find the goal node's parent, then the parent's parent, and so on, until we reach the start state, and then we revert this path. What is the path found by DFS for this problem?

_Points:_ 1.5

```python
# Express your answer as a list of tuples starting from the start state and ending in the goal state. If no plan is found, set the plan to []
plan_124 = ...
```

#### 1.2.5)

What is the total cost of the plan found by A*?

_Points:_ 1

```python
# Express your answer as a list of tuples starting from the start state and ending in the goal state. If no plan is found, set the plan to []
ans_125 = ...
```

#### 1.2.6)

What can you say about the plan found by A* using the octile heuristic?

a. It is the optimal plan, because A* is always guaranteed to find the optimal plan\
b. It is not optimal because the heuristic is not admissible\
c. It is not optimal because the heuristic is not consistent\
d. It is not optimal even though the heuristic is admissible\
e. It is the optimal plan even though the heuristic is not admissible

_Points:_ 3

```python
# Enter your answer as a lower-case Python string
ans_126 = ...
```

### 1.3) A* with Conveyor-Belt Heuristic (19 pts)

Now suppose that we wish to incorporate knowledge about the conveyor belts in our heuristic. Our conveyor-belt heuristic assumes that all horizontal and vertical actions have a cost of 0.2 (so there is a conveyor belt in every one of those edges) and diagonal actions have a cost of 1.4. 

Which of the following formulas implements this heuristic function?

a. $\max (\Delta x, \Delta y) + 0.2  \cdot \min(\Delta x, \Delta y)$\
b. $0.2 \cdot (\Delta x + \Delta y) + 0.4 \cdot \min (\Delta x, \Delta y)$\
c. $0.2 (\Delta x + \Delta y)$\
d. $0.2 \cdot (\Delta x + \Delta y) + 1.2 \cdot \min (\Delta x, \Delta y)$

_Points:_ 3

```python
# Enter your answer as a lower-case Python string (e.g., "a" for option a) 
ans_13 = ...
```

Let's now trace the A* algorithm on this problem with the conveyor-belt heuristic. The initial priority queue and parent dictionary are:
```
q = {(0,0): 0.8}
p = {(0,0): None}
```


#### 1.3.1)

In the first step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the first iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_131 = ...
```

After the first step, what are the contents of the queue (i.e., after extracting the first node `v_131` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_131 = ...
```

After the first step, what are the contents of the parents directory (i.e., after extracting the first node `v_121` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_121 = ...
```

#### 1.3.2)

In the second step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the second iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_122 = ...
```

After the second step, what are the contents of the queue (i.e., after extracting the second node `v_122` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_132 = ...
```

After the second step, what are the contents of the parents directory (i.e., after extracting the second node `v_122` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_132 = ...
```

#### 1.3.3)

In the third step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the third iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_133 = ...
```

After the third step, what are the contents of the queue (i.e., after extracting the third node `v_133` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_133 = ...
```

After the third step, what are the contents of the parents directory (i.e., after extracting the third node `v_133` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_133 = ...
```

#### 1.3.4)

In the fourth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the fourth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_134 = ...
```

After the fourth step, what are the contents of the queue (i.e., after extracting the fourth node `v_134` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_134 = ...
```

After the fourth step, what are the contents of the parents directory (i.e., after extracting the fourth node `v_134` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_134 = ...
```

#### 1.3.5)

In the fifth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the fifth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_135 = ...
```

After the fifth step, what are the contents of the queue (i.e., after extracting the fifth node `v_135` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_135 = ...
```

After the fifth step, what are the contents of the parents directory (i.e., after extracting the fifth node `v_135` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_135 = ...
```

#### 1.3.6)

In the sixth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the sixth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_136 = ...
```

After the sixth step, what are the contents of the queue (i.e., after extracting the sixth node `v_136` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_136 = ...
```

After the sixth step, what are the contents of the parents directory (i.e., after extracting the sixth node `v_136` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_136 = ...
```

#### 1.3.7)

In the seventh step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the seventh iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_137 = ...
```

After the seventh step, what are the contents of the queue (i.e., after extracting the seventh node `v_137` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_137 = ...
```

After the seventh step, what are the contents of the parents directory (i.e., after extracting the seventh node `v_137` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_137 = ...
```

#### 1.3.8)

In the eighth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the eighth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_138 = ...
```

After the eighth step, what are the contents of the queue (i.e., after extracting the eighth node `v_138` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_138 = ...
```

After the eighth step, what are the contents of the parents directory (i.e., after extracting the eighth node `v_138` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_138 = ...
```

#### 1.3.9)

In the ninth step, what is the current node inspected by the algorithm (i.e., the node extracted from the queue `q` at the ninth iteration)?

_Points:_ 0.5

```python
# Express your answer as a tuple representing the node (e.g., (2,3)). If the algorithm has terminated, set the value to None.
v_139 = ...
```

After the ninth step, what are the contents of the queue (i.e., after extracting the ninth node `v_139` and adding its neighbors to the queue, what is the current state of the queue `q`)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing nodes) to
# scalar values (representing costs from start to node) (e.g., {(0,0): 0}). If 
# the algorithm has terminated, set the value to an empty dictionary {}
q_139 = ...
```

After the ninth step, what are the contents of the parents directory (i.e., after extracting the ninth node `v_139` and adding its neighbors to the queue, what are each visited node's parents)? 

_Points:_ 0.5

```python
# Express your answer as a dictionary mapping tuples (representing child nodes) to
# tuples (representing parent nodes) (e.g., {(0,0): None, (2,1): (a,b)} where (a,b) in 
# this example is the parent of (2,1)). If the algorithm has terminated, set the value
# to the latest parents dictionary
p_139 = ...
```

#### 1.3.10)

When A* succeeds, we can find a plan by tracing back the parents in the `p` dictionary. In particular, we find the goal node's parent, then the parent's parent, and so on, until we reach the start state, and then we revert this path. What is the path found by A* for this problem?

_Points:_ 1.5

```python
# Express your answer as a list of tuples starting from the start state and ending in the goal state. If no plan is found, set the plan to []
plan_1310 = ...
```

#### 1.3.11)

What is the total cost of the plan found by A*?

_Points:_ 1

```python
# Express your answer as a scalar
ans_1311 = ...
```

## 1.4) Admissibility and consistency

Recall our definitions of admissibility and consistency of heuristic functions. A heuristic function $h$ is admissible if it never overestimates the true cost $c$ to reach the goal from a node: $h(v) \leq c(v) \quad \forall v \in \mathcal{V}$. A heuristic function $h$ is consistent if it satisfies the triangle inequlaity: the heuristic of a node is less than or equal to the cost to move to a neighbor plus the heuristic of that neighbor, $h(v) \leq c(v, v') + h(v') \quad \forall v\in\mathcal{V}, v'\in N(v)$. Notably, any consistent heuristic is by definition admissible (but not vice versa). 

For each of the following heuristic functions in our conveyor-belt problem, determine whether the heuristic is:
a. Admissible but not consistent\
b. Consistent and admissible\
c. Not admissible

#### 1.4.1)

The all-zero heuristic, given by $h(x,y) = 0$

_Points:_ 2

```python
# Enter your answer as a lower-case Python string (e.g., "a")
ans_141 = ...
```

#### 1.4.2)

Euclidean distance, given by $h(x,y) = \sqrt{(x-2)^2 + (y-2)^2}$. You can use $\sqrt{2}\approx 1.4$.

_Points:_ 2

```python
# Enter your answer as a lower-case Python string (e.g., "a")
ans_142 = ...
```

#### 1.4.3)

Euclidan distance scaled by 0.2, given by $h(x,y) = 0.2\cdot\sqrt{x^2 + y^2}$. You can use \sqrt{2}\approx 1.4 and $\sqrt{5}\approx2.25.

_Points:_ 2

```python
# Enter your answer as a lower-case Python string (e.g., "a")
ans_143 = ...
```

#### 1.4.4)

Manhattan distance minus 1.6, clipped at 0: $h(x,y) = \max(0, (2-x) + (2-y) - 1.6).

_Points:_ 2

```python
# Enter your answer as a lower-case Python string (e.g., "a")
ans_144 = ...
```

#### 1.4.5)

Manhattan distance scaled by 0.2 (or our conveyor-belt heuristic): $h(x,y) = 0.2 \cdot ( 2-x + 2 -y).

_Points:_ 2

```python
# Enter your answer as a lower-case Python string (e.g., "a")
ans_145 = ...
```

#### 1.4.6)

Octile distance, defined in Problem 1.2.

_Points:_ 2

```python
# Enter your answer as a lower-case Python string (e.g., "a")
ans_146 = ...
```

## 2) Implementing A* for robot planning (60 pts)

We're now going to write a Python implementation of A*. 



### 2.1) Priority queue as an open set

While the open sets we used last week in HW 1 were restricted to the two main operations (`add` and `next`), A* requires the addition of a `decrease_priority` operation. 

We will write a Python class that implements a priority queue open set. The combination of `add`, `extract_min`, and `decrease_priority` should be such that the node with the least priorty is returned each time. To simplify our implementation, we will use the [`heapdict`](https://github.com/DanielStutzbach/heapdict) (you can `pip install heapdict`), which allows us to add and modify queue elements as we do with dictionaries (e.g., `q[node] = val`) and offers the `popitem()` method to remove the lowest priority element in the queue. 

I have not added public tests for this class, but you are welcome to add some of your own.

*Note: the `heapdict` implementation does not return elements in the order they are inserted, for various reasons. You **do not** need to write code to recover this order.*

_Points:_ 5

```python
from heapdict import heapdict

class PriorityQueue_21():
    def __init__(self):
        self.q = ...

    def add(self, node, value):

    def extract_min(self):
        ...

    def decrease_priority(self, node, value):

    def __len__(self):
        """ Enables calling len(obj) on objects of the class"""
        return len(self.q)

    def __contains__(self, node):
        """ Enables calling `if node in obj` on objects of the class"""
        return node in self.q
```

### 2.2) Costs and heuristics for occupancy grids

Like in HW 1, we will represent 2D robotic navigation problems via an occupancy grid, as exemplified below:

```
[[0, 0, 1, 0, 0],
 [0, 1, 0, 0, 0],
 [0, 0, 0, 1, 0],
 [1, 0, 1, 0, 0],
 [0, 0, 0, 0, 0]]
```

This time, we will consider an eight-connected grid, with actions: up, down, left, right, up-left, up-right, down-left, and down-right. Vertical and horizontal actions have a cost of 1 and diagonal actions have a cost of 1.4. 


### 2.2.1) Cost function

Given pair of neighbors `u` and `v` in the grid, write a Python function that returns the cost of going from `u` to `v`. `u` and `v` are both expressed as python tuples.

_Points:_ 2.5

```python
def eight_connected_cost_221(u, v):
    ...
```

### 2.2.2) Manhattan heuristic

Given a node `u` and a goal node `g`, write a function that implements the Manhattan distance heuristic. 

_Points:_ 2.5

```python
def manhattan_heuristic_222(u, g):
    ...
```

### 2.2.3) Octile heuristic

Given a node `u` and a goal node `g`, write a function that implements the octile distance heuristic. 

_Points:_ 2.5

```python
def octile_heuristic_223(u, g):
    ...
```

### 2.2.4) Dijkstra's heuristic

Given a node `u` and a goal node `g`, write a function that implements a heuristic function such that A* using this heuristic is equivalent to Dijkstra's algorithm.

_Points:_ 2.5

```python
def dijkstra_heuristic_224(u, g):
    ...
```

#### 2.3) A* search on grids

We are now ready to write our A* grid search function, `astar_grid_search_23`. The function takes the following inputs:
- `grid`: a nested 2D Python array of 1s and 0s representing an occupancy grid
- `start`: the `(i,j)` coordinates of the start node
- `goal`: the `(i,j)` coordinates of the goal node
- `open_set`: an object with methods `add`, `extract_min`, `decrease_priority`, and `__len__`
- `cost`: a function that computes the cost to go from one node `u` to another node `v`
- `heuristic`: a function that computes the heuristic estimate of the cost to go from one node `u` to a goal node `g`

The function will return two values:
- `plan`: a list of nodes that represents a valid plan starting from `start` and ending in `goal`, or an empty list if no plan is found. You might find it useful to include your `retrace_plan` function from HW 1.
- `node_history`: a list of nodes in the order they were expanded by the algorithm (i.e., the nodes that come from `extract_min`)

When adding neighbors into the open set, follow the convention from the previous problem, which orders nodes first according to their `i` coordinate and then according to their `j` coordinate. *Note: because `heapdict` does not maintain this order, the order in which nodes are visited is not the same as the order in which they are added. This is okay: our convention for this problem will be that the order is that in which `heapdict` returns elements after we've added them in sorted order.*

_Points:_ 15

```python
from collections import defaultdict

def astar_grid_search_23(grid, start, goal, open_set, cost, heuristic):
    visited_nodes = []
    c2come = defaultdict(lambda: float('inf'))
    ctot = defaultdict(lambda: float('inf'))
    parent = {}

    h = heuristic(start, goal)
    c2come[start] = 0
    ctot[start] = h
    open_set.add(start, h)
    parent[start] = None
    while len(open_set) > 0:
        v = ...
        visited_nodes.append(v)
        if v == goal:
            return ...
        # Loop over neighbors of v, check their validity, and update the open_set using cost and heuristic
        # *** YOUR CODE HERE ***
    # If no plan is found, return an empty plan
    return [], visited_nodes
```

#### 2.3.1) A* and Dijkstra

Put everything together to implement the A* (using the octile heuristic) and Dijkstra algorithms. Use eight-connected cost for both. *Hint: these should be very simple one-line functions.*

To help with debugging, I have included tests for both algorithms using the following grid and testing for plans from `(0,0)` to `(4,4)`:

```
[[0, 0, 1, 0, 0],
 [0, 1, 0, 0, 0],
 [0, 0, 0, 1, 0],
 [1, 0, 1, 0, 0],
 [0, 0, 0, 0, 0]]
```

_Points:_ 10

```python
def astar_231(grid, start, goal):
    ...

def dijkstra_231(grid, start, goal):
    ...
```

```python
grader.check("q2.3.1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Fill out the answers to all questions and submit your file hw2.ipynb to the HW2 assignment on Gradescope. You are free to resubmit as many times as you wish.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.1.1.1 (points: 0.5)


Tests for q1.1.1.2 (points: 0.5)


Tests for q1.1.1.3 (points: 0.5)


Tests for q1.1.10 (points: 1.5)


Tests for q1.1.11 (points: 1)


Tests for q1.1.2.1 (points: 0.5)


Tests for q1.1.2.2 (points: 0.5)


Tests for q1.1.2.3 (points: 0.5)


Tests for q1.1.3.1 (points: 0.5)


Tests for q1.1.3.2 (points: 0.5)


Tests for q1.1.3.3 (points: 0.5)


Tests for q1.1.4.1 (points: 0.5)


Tests for q1.1.4.2 (points: 0.5)


Tests for q1.1.4.3 (points: 0.5)


Tests for q1.1.5.1 (points: 0.5)


Tests for q1.1.5.2 (points: 0.5)


Tests for q1.1.5.3 (points: 0.5)


Tests for q1.1.6.1 (points: 0.5)


Tests for q1.1.6.2 (points: 0.5)


Tests for q1.1.6.3 (points: 0.5)


Tests for q1.1.7.1 (points: 0.5)


Tests for q1.1.7.2 (points: 0.5)


Tests for q1.1.7.3 (points: 0.5)


Tests for q1.1.8.1 (points: 0.5)


Tests for q1.1.8.2 (points: 0.5)


Tests for q1.1.8.3 (points: 0.5)


Tests for q1.1.9.1 (points: 0.5)


Tests for q1.1.9.2 (points: 0.5)


Tests for q1.1.9.3 (points: 0.5)


Tests for q1.2 (points: 3)


Tests for q1.2.1.1 (points: 0.5)


Tests for q1.2.1.2 (points: 0.5)


Tests for q1.2.1.3 (points: 0.5)


Tests for q1.2.2.1 (points: 0.5)


Tests for q1.2.2.2 (points: 0.5)


Tests for q1.2.2.3 (points: 0.5)


Tests for q1.2.3.1 (points: 0.5)


Tests for q1.2.3.2 (points: 0.5)


Tests for q1.2.3.3 (points: 0.5)


Tests for q1.2.4 (points: 1.5)


Tests for q1.2.5 (points: 1)


Tests for q1.2.6 (points: 3)


Tests for q1.3 (points: 3)


Tests for q1.3.1.1 (points: 0.5)


Tests for q1.3.1.2 (points: 0.5)


Tests for q1.3.1.3 (points: 0.5)


Tests for q1.3.10 (points: 1.5)


Tests for q1.3.11 (points: 1)


Tests for q1.3.2.1 (points: 0.5)


Tests for q1.3.2.2 (points: 0.5)


Tests for q1.3.2.3 (points: 0.5)


Tests for q1.3.3.1 (points: 0.5)


Tests for q1.3.3.2 (points: 0.5)


Tests for q1.3.3.3 (points: 0.5)


Tests for q1.3.4.1 (points: 0.5)


Tests for q1.3.4.2 (points: 0.5)


Tests for q1.3.4.3 (points: 0.5)


Tests for q1.3.5.1 (points: 0.5)


Tests for q1.3.5.2 (points: 0.5)


Tests for q1.3.5.3 (points: 0.5)


Tests for q1.3.6.1 (points: 0.5)


Tests for q1.3.6.2 (points: 0.5)


Tests for q1.3.6.3 (points: 0.5)


Tests for q1.3.7.1 (points: 0.5)


Tests for q1.3.7.2 (points: 0.5)


Tests for q1.3.7.3 (points: 0.5)


Tests for q1.3.8.1 (points: 0.5)


Tests for q1.3.8.2 (points: 0.5)


Tests for q1.3.8.3 (points: 0.5)


Tests for q1.3.9.1 (points: 0.5)


Tests for q1.3.9.2 (points: 0.5)


Tests for q1.3.9.3 (points: 0.5)


Tests for q1.4.1 (points: 2)


Tests for q1.4.2 (points: 2)


Tests for q1.4.3 (points: 2)


Tests for q1.4.4 (points: 2)


Tests for q1.4.5 (points: 2)


Tests for q1.4.6 (points: 2)


Tests for q2.1 (points: 5)


Tests for q2.2.1 (points: 2.5)


Tests for q2.2.2 (points: 2.5)


Tests for q2.2.3 (points: 2.5)


Tests for q2.2.4 (points: 2.5)


Tests for q2.3 (points: 15)


Tests for q2.3.1 (points: 10)
```
>>> grid = [[0, 0, 1, 0, 0], [0, 1, 0, 0, 0], [0, 0, 0, 1, 0], [1, 0, 1, 0, 0], [0, 0, 0, 0, 0]]
>>> start = (0, 0)
>>> goal = (4, 4)
>>> plan, visited = astar_231(grid, start, goal)
>>> plan == [(0, 0), (1, 0), (2, 1), (2, 2), (3, 3), (4, 4)] and visited == [(0, 0), (1, 0), (2, 1), (2, 2), (3, 3), (4, 4)]
True
```
```
>>> grid = [[0, 0, 1, 0, 0], [0, 1, 0, 0, 0], [0, 0, 0, 1, 0], [1, 0, 1, 0, 0], [0, 0, 0, 0, 0]]
>>> start = (0, 0)
>>> goal = (4, 4)
>>> plan, visited = dijkstra_231(grid, start, goal)
>>> plan == [(0, 0), (0, 1), (1, 2), (2, 2), (3, 3), (4, 4)] and visited == [(0, 0), (1, 0), (0, 1), (2, 0), (1, 2), (2, 1), (1, 3), (3, 1), (2, 2), (0, 3), (1, 4), (4, 1), (4, 0), (2, 4), (3, 3), (4, 2), (0, 4), (3, 4), (4, 3), (4, 4)]
True
```
